# CSCI 3155 Spring 2023
## Recitation Week 14: Traits and Generics

## Traits

Traits are an important mechanism for code reuse in scala. They allow us to define functionality that can be 
exported across multiple objects in the overall hierarchy. A trait is almost like an abstract class or an interface. It can define its own members and methods.

For this exercise, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.  We also have an abstract class `Animal` as a superclass for all animals.

```scala
abstract class Animal

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined [32mclass[39m [36mAnimal[39m
defined [32mtrait[39m [36mNumberOfLegs[39m
defined [32mtrait[39m [36mWarmBlooded[39m
```

### Example
We will define a class `Human` which will extend the appropriate classes and traits defined above.
Our class will also take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

Which traits should we use for our human class? (Discuss this first then move onto the coding of the class)

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Human(val name: String) extends WarmBlooded with NumberOfLegs {
    override val nLegs:Int = 2
    val bodyTempMaintained = 98.0
    def getName: String = name
    // below is also valid, but may change testing needs
    // def getName(): String = name
}
//END SOLUTION
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val t1 = new Human("Jane Smith")
t1.isInstanceOf[Animal]
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
//END TEST
```

```scala
// ALTERNATE SOLUTION
class Human(val name: String, 
            val nLegs: Int = 2, 
            val bodyTempMaintained: Double = 98.0) extends WarmBlooded 
        with NumberOfLegs { 
    def getName(): String = this.name
}
val t1 = new Human("Jane Smith")
t1.isInstanceOf[Animal]
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
```

Output:
```
defined [32mclass[39m [36mHuman[39m
[36mt1[39m: [32mHuman[39m = ammonite.$sess.cmd1$Helper$Human@52d2fec6
[36mres1_2[39m: [32mBoolean[39m = [32mtrue[39m
```

### Exercise
Now we will define a class named `Table`.

Which of the traits and classes should we extend? (discuss this first then advance to the coding) (if someone complains about the number of legs on a table, adjust the exercise so that the table class takes in the number of legs, and update the test)

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```scala
//BEGIN SOLUTION
class Table extends NumberOfLegs {
    val nLegs = 4
}
//END SOLUTION
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
//END TEST
```

```scala
// ALTERNATE SOLUTION
class Table(val nLegs: Int = 4) extends NumberOfLegs
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
```

Output:
```
defined [32mclass[39m [36mTable[39m
[36mtbl[39m: [32mTable[39m = ammonite.$sess.cmd2$Helper$Table@6ada9223
```

## Type Casting
We may want to do certain things based on a more precise type than a trait gives us, like only the humans in a list of instances of NumberOfLegs.
While not as easy as pattern matching with case classes, this is possible by using the `isInstanceOf[T]` and the `asInstanceOf[T]` functions.

The expression `obj.isInstanceOf[T]` returns true if `obj` is an instance of `T` and false otherwise.
Similarly, `obj.asInstanceOf[T]` returns the same instance `obj`, but with the given type `T` (if it is an instance of `T`).

### Demo
Consider the following code and it's errors

```scala
// // fails 
// val a0: NumberOfLegs = new Human("Liskov")
// a0.getName
```

```scala
// works
val a0: NumberOfLegs = new Human("Liskov")
assert(a0.isInstanceOf[Human])
val h0: Human = a0.asInstanceOf[Human]
val name: String = h0.getName
```

```scala
// // fails
// val a1: NumberOfLegs = new Table()
// assert(a1.isInstanceOf[Human])
```

### Exercise
To try out these functions, lets write a function to get the names of all the humans in a list of instances of the NumberOfLegs trait.

```scala
def getNames(leggedThings : List[NumberOfLegs]): List[String] = {
    leggedThings match {
        case h :: t =>
            // begin solution
            if (h.isInstanceOf[Human]) 
                (h.asInstanceOf[Human].getName):: getNames(t)
            else 
                getNames(t)
            // end solution
        case Nil => Nil
    }
}
```

```scala
// TEST CASE
val leggedThings : List[NumberOfLegs] = List(new Human("Charles"), new Human("Sukanya"), new Table(), new Table(), new Human("Abhishek"))
assert(getNames(leggedThings) == List("Charles", "Sukanya", "Abhishek"), "Names don't match")
```

```scala
// Challenge option: reimplement with functors
def getNames(leggedThings : List[NumberOfLegs]): List[String] = {
    // begin solution
    leggedThings filter {
        leggedThing => leggedThing.isInstanceOf[Human]
    } map {
        leggedThing => leggedThing.asInstanceOf[Human]
    } map {
        human => human.getName
    } 
    // end solution
}
```

Output:
```
defined [32mfunction[39m [36mgetNames[39m
```

```scala
// TEST CASE
val leggedThings : List[NumberOfLegs] = List(new Human("Charles"), new Human("Sukanya"), new Table(), new Table(), new Human("Abhishek"))
assert(getNames(leggedThings) == List("Charles", "Sukanya", "Abhishek"), "Names don't match")
```

Output:
```
[36mleggedThings[39m: [32mList[39m[[32mNumberOfLegs[39m] = [33mList[39m(
  ammonite.$sess.cmd1$Helper$Human@64995226,
  ammonite.$sess.cmd1$Helper$Human@674fb721,
  ammonite.$sess.cmd2$Helper$Table@4e29da44,
  ammonite.$sess.cmd2$Helper$Table@1564ff3b,
  ammonite.$sess.cmd1$Helper$Human@62f73263
)
```

```scala
// Challenge option: reimplement with functors
// ALTERNATIVE SOLUTION
def getNames(leggedThings : List[NumberOfLegs]): List[String] = {
    // begin solution
    (leggedThings foldRight[List[String]] (Nil)){
        case (h, acc) if h.isInstanceOf[Human] => 
                h.asInstanceOf[Human].getName :: acc
        case (_, acc) => acc
    } 
    // end solution
}
// TEST CASE
val leggedThings : List[NumberOfLegs] = List(new Human("Charles"), new Human("Sukanya"), new Table(), new Table(), new Human("Abhishek"))
assert(getNames(leggedThings) == List("Charles", "Sukanya", "Abhishek"), "Names don't match")
```

Output:
```
defined [32mfunction[39m [36mgetNames[39m
[36mleggedThings[39m: [32mList[39m[[32mNumberOfLegs[39m] = [33mList[39m(
  ammonite.$sess.cmd1$Helper$Human@1a30e30d,
  ammonite.$sess.cmd1$Helper$Human@1d054b03,
  ammonite.$sess.cmd2$Helper$Table@760b482c,
  ammonite.$sess.cmd2$Helper$Table@6b3ee56b,
  ammonite.$sess.cmd1$Helper$Human@58464477
)
```

## Generics

Generic classes and functions are classes that take a type as a parameter.
We put these type parameters in square brackets (for example `List[String]` or `asInstanceOf[Human]` where we are passing `String` or `Human`, respectively)

One common use case for generics is for collections like maps, trees, and lists.
We will further exam generics using the example of lists.

If we wanted to define a list of integers, we could do something like this:

```scala
sealed trait IntList
case class IntNil() extends IntList
case class IntCons(head: Int, tail: IntList) extends IntList
```

But what if we needed a list of strings instead?
We could write the following:

```scala
sealed trait StringList
case class StringNil() extends StringList
case class StringCons(head: String, tail: StringList) extends StringList
```

What about booleans? Or your custom class you wrote, such as `Human` above?
We don't want to redefine these traits and the functions that work with them everytime we need a list for a different type.
We could define a list holding objects of type `Any` (the supertype of all types), but this has problems when we try access the objects we put in the list.

Instead we use generics.
We do this by adding type parameters to our code.

Coming back to our list example, we could write it using generics as follows:

```scala
sealed trait GenericList[T]

case class GenericNil[T]() extends GenericList[T]
case class GenericCons[T](h: T, t: GenericList[T]) extends GenericList[T]
```

The type parameter `[T]` can be replaced with any class we want, either explicitly or by the type scala infers:

```scala
val explicit = GenericCons[String]("hello", GenericCons[String]("world", GenericNil[String]()))
val inferred = GenericCons("hello", GenericCons("world", GenericNil()))
```

```scala
val l = GenericCons[Integer](1, GenericNil())
```

### Exercise
As an exercise with our `GenericList`, let's define the map function for `GenericList`s

```scala
def map[A,B](list: GenericList[A], function: (A) => B): GenericList[B] = {
    // begin solution
    list match {
        case GenericCons(h, t) => GenericCons[B](function(h), map(t, function))
        case GenericNil() => GenericNil[B]()
    }
    // end solution
}
```

```scala
val newL = map[Integer, String](l, x => s"${x}")
```

```scala
newL
```

### Type Bounds (If time permits)
We can also apply restrictions on the the types used for our type parameters

We can use `<:` to ensure that the type parameter is a subtype of the given type, or `:>` to ensure that the type parameter is a supertype of the given type.

```scala
sealed trait LegList[T <: NumberOfLegs]

case class LegNil[T <: NumberOfLegs]() extends LegList[T]
case class LegCons[T <: NumberOfLegs](h: T, t: LegList[T]) extends LegList[T]
```

```scala
def countLegs[T <: NumberOfLegs](leggedThings : LegList[T]) : Int = leggedThings match {
    case LegNil() => 0
    case LegCons(h, t) => h.getNumberOfLegs + countLegs(t)
}
```

```scala
def countLegs[T <: NumberOfLegs](leggedThings : List[T]): Int = {
    leggedThings.foldLeft[Int](0)((accum, leggedThing) => accum + leggedThing.getNumberOfLegs)
}
```

Note that the ever useful `foldLeft` function also has a type parameter, which lets `foldLeft` return any type we want.  Of course, the compiler usually infers the type parameter's value, so we don't need include the type.

## That's all folks!